## Q2 (2 points)
- เขียนฟังก์ชัน get_dict(file_name) โดยรับพารามิเตอร์เป็น file_name ที่สร้าง dictionary จากการอ่านไฟล์ conversation.txt (มีไฟล์ให้)  <br>
    - จากนั้นนับความถี่ของแต่ละคำที่ปรากฏในไฟล์ โดยแต่ละคำคั่นด้วยช่องว่าง และให้ตัด **: ? , .** ที่อยู่หน้าหรือหลังคำ 
ออกจากข้อความด้วย <br> และแสดงผลลัพธ์เป็นข้อมูลใน dictionary ตามตัวอย่างการทำงาน 

### ตัวอย่างการทำงาน: 
get_dict(**'conversation1.txt'**) <br>
{'Tom': 2, 'Hi': 1, 'I': 2, 'am': 2, 'Ana': 2}

### ตัวอย่างการทำงาน: 
get_dict(**'conversation2.txt'**) <br>
{'Mr': 3, 'Bean': 3, 'Good': 2, 'morning': 2, 'Mrs': 2, 'Jones': 2, 'How': 1, 'are': 1, 'you': 1, "I'm": 1, 'fine': 1, 'Thanks': 1}

### ตัวอย่างการทำงาน: 
get_dict(**'conversation3.txt'**) <br>
{'Tom': 3, "It's": 2, 'an': 1, 'ugly': 1, 'day': 1, 'today': 2, 'Jerry': 2, 'I': 2, 'know': 1, 'think': 1, 'it': 2, 'may': 1, 'rain': 2, 'the': 1, 'middle': 1, 'of': 1, 'summer': 1, "shouldn't": 1, 'That': 1, 'would': 1, 'be': 1, 'weird': 1, 'Yeah': 1, 'especially': 1, 'since': 1, "it's": 1, 'ninety': 1, 'degrees': 1, 'outside': 1}

<!-- grade_id: name, points: 0, grade: false, solution: true -->
```python
def get_dict(file):
    sol = open(file, 'r')
    score = sol.readlines()
    sol.close()
    word = []
    for i in score :
        cn = i.split()
        for j in cn :
            ct = j.strip(':?,.')
            word.append(ct)
    value = {}
    for i in word :
        if i not in value :
            value[i] = 1
        else :
            value[i] = value[i] + 1
    return value
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
'''ตัวอย่าง get_dict("ใส่ชื่อไฟล์.txt")'''
get_dict('conversation2.txt')
```

Output:
```
{'Mr': 3,
 'Bean': 3,
 'Good': 2,
 'morning': 2,
 'Mrs': 2,
 'Jones': 2,
 'How': 1,
 'are': 1,
 'you': 1,
 "I'm": 1,
 'fine': 1,
 'Thanks': 1}
```

<!-- grade_id: cell-b082fbfab4c14322, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
assert get_dict('conversation1.txt') == {'Tom': 2, 'Hi': 1, 'I': 2, 'am': 2, 'Ana': 2}
assert get_dict('conversation2.txt') == {'Mr': 3, 'Bean': 3, 'Good': 2, 'morning': 2, 'Mrs': 2, 'Jones': 2, 'How': 1, 'are': 1, 'you': 1, "I'm": 1, 'fine': 1, 'Thanks': 1}
print('OK')
```

Output:
```
Cell In[14], line 2
      1 '''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
----> 2 assert get_dict('conversation1.txt') == {'Tom': 2, 'Hi': 1, 'I': 2, 'am': 2, 'Ana': 2}
      3 assert get_dict('conversation2.txt') == {'Mr': 3, 'Bean': 3, 'Good': 2, 'morning': 2, 'Mrs': 2, 'Jones': 2, 'How': 1, 'are': 1, 'you': 1, "I'm": 1, 'fine': 1, 'Thanks': 1}
      4 print('OK')

Cell In[12], line 13, in get_dict(file)
     11 value = {}
     12 for i in word :
---> 13     if i not in count :
     14         value[i] = 1
     15     else :

NameError: name 'count' is not defined
```

<!-- grade_id: cell-4edb7a354d38a7f0, points: 1, grade: true, solution: false -->
```python
'''เทสลับ'''
# ### BEGIN HIDDEN TESTS
assert get_dict('conversation3.txt') == {'Tom': 3, "It's": 2, 'an': 1, 'ugly': 1, 'day': 1, 'today': 2, 'Jerry': 2, 'I': 2, 'know': 1, 'think': 1, 'it': 2, 'may': 1, 'rain': 2, 'the': 1, 'middle': 1, 'of': 1, 'summer': 1, "shouldn't": 1, 'That': 1, 'would': 1, 'be': 1, 'weird': 1, 'Yeah': 1, 'especially': 1, 'since': 1, "it's": 1, 'ninety': 1, 'degrees': 1, 'outside': 1},get_dict('conversation3.txt')
assert get_dict('conversation4.txt') == {'Baby': 3, 'Shark': 3, 'doodoo': 6},get_dict('conversation4.txt')
# ### END HIDDEN TESTS
```

```python

```

```python

```