## Q3 (3 points)
จงเขียนโปรแกรมเพื่อตัดสินผลการประกวดร้องเพลง โดยให้โปรแกรมอ่านคะแนนที่ผู้เข้าประกวดได้รับจากคณะกรรมการ <br>
(เต็ม 10 คะแนนสำหรับกรรมการแต่ละท่าน) และผู้ชมในห้องส่ง (เต็ม 70 คะแนน) <br>
จากไฟล์ score\<n\>.txt (แต่ละบรรทัดประกอบด้วยรหัสผู้ประกวดร้องเพลง คะแนนจากกรรมการคนที่ 1 คะแนนจากกรรมการคนที่ 2 คะแนนจากกรรมการคนที่ 3 และคะแนนจากผู้ชมในห้องส่ง) <br>

กำหนดให้เขียนฟังก์ชัน `get_dict(file_name)` ที่มี parameter เป็นชื่อ file แล้วสร้างและคืนค่า dictionary จากการอ่านไฟล์ โดยมี
- key เป็นสตริงของรหัสผู้เข้าประกวด 
- value เป็นผลรวมคะแนนทั้งหมด

จงเขียนฟังก์ชัน `find_top3(scores)` ที่มี parameter เป็น dictionary ของผู้เข้าประกวดและคะแนน แล้วคืนค่าเป็น list ของผู้มีคะแนนสูงสุดการประกวด 3 อันดับแรก กำหนดให้ไม่มีผู้ใดได้คะแนนเท่ากัน

### ตัวอย่างไฟล์  score1.txt: 

CU01 8 8 9 65 <br>
CU02 6 6 6 45 <br>
CU03 8 8 7 55 <br>
CU04 6 8 6 57 <br>
CU05 9 9 9 66 <br>

### ตัวอย่างไฟล์ score2.txt: 

CU01 8 8 9 65 <br>
CU02 6 6 6 45 <br>
CU03 8 7 7 55 <br>
CU04 6 8 6 57 <br>
CU05 9 9 9 66 <br>
CU06 8 8 8 62 <br>
CU07 10 8 9 67 <br>

### ตัวอย่างไฟล์ score3.txt: 

CU01 8 8 9 65 <br>
CU02 6 6 6 45 <br>
CU03 8 7 7 55 <br>
CU04 6 8 6 57 <br>
CU05 9 9 9 66 <br>
CU06 8 8 8 62 <br>
CU07 10 8 9 67 <br>
CU08 8 7 9 60 <br>
CU09 7 7 7 58 <br>
CU10 9 8 9 60 <br>

### ตัวอย่างการทำงาน: 

Choose your score: **score1.txt**<br>
The winners are :<br>
['CU05', 'CU01', 'CU03']

### ตัวอย่างการทำงาน: 

Choose your score: **score2.txt**<br>
The winners are :<br>
['CU07', 'CU05', 'CU01']

### ตัวอย่างการทำงาน: 

Choose your score: **score3.txt**<br>
The winners are :<br>
['CU07', 'CU01', 'CU05']

<!-- grade_id: name, points: 0, grade: false, solution: true -->
```python
def get_dict(file_name):
    st = open(file_name)
    st = st.read().splitlines()
    r = {}
    for i in st :
        i = i.split()
        s1 = int(i[1])
        s2 = int(i[2])
        s3 = int(i[3])
        s4 = int(i[4])
        sum = s1+s2+s3+s4
        r[i[0]] = sum
    return r

def find_top3(scores):
    '''คืน list ของผู้มีคะแนนสูงสุดการประกวด 3 อันดับแรก'''
    sd = sorted(scores.items(),key=lambda x:x[1],reverse=True)
    t3 = [a for a,b in sd]
    t3 = t3[0:3]
    return t3

def run():
    '''รับชื่อไฟล์คะแนน แล้วแสดง list ของผู้ชนะ 3 ลำดับแรก'''
    x = input("Input file_name :")
    scores = get_dict(x)
    print('The winners are :')
    print(find_top3(scores))
```

```python
'''ตรวจสอบการสร้าง dict'''
'''ตัวอย่าง get_dict("ใส่ชื่อไฟล์.txt")'''
get_dict('score2.txt')
```

Output:
```
{'CU01': 90,
 'CU02': 63,
 'CU03': 77,
 'CU04': 77,
 'CU05': 93,
 'CU06': 86,
 'CU07': 94}
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ทดสอบฟังก์ชัน get_dict(file_name)'''
assert get_dict('score1.txt') == {'CU01': 90, 'CU02': 63, 'CU03': 78, 'CU04': 77, 'CU05': 93}
assert get_dict('score2.txt') == {'CU01': 90, 'CU02': 63, 'CU03': 77, 'CU04': 77, 'CU05': 93, 'CU06': 86, 'CU07': 94}
assert get_dict('score3.txt') == {'CU01': 90, 'CU02': 63, 'CU03': 77, 'CU04': 77, 'CU05': 87, 'CU06': 86, 'CU07': 94, 'CU08': 84, 'CU09': 79, 'CU10': 86} 
print('OK')
```

Output:
```
OK
```

<!-- grade_id: cell-a90ea8f11e54426e, points: 1, grade: true, solution: false -->
```python
'''ทดสอบฟังก์ชัน find_top3()'''
'''จากตัวอย่าง dictionary ผู้มีคะแนนสูงสุด 3 ลำดับแรกคือ E, B และ F ตามลำดับ'''
assert find_top3({'A':80, 'B':90, 'C':70, 'D':50, 'E':95, 'F':86}) == ['E', 'B', 'F']
assert find_top3({'A':97, 'B':68, 'C':96, 'D':88, 'E':85, 'F':86}) == ['A', 'C', 'D']
print('OK')
```

Output:
```
OK
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
Input file_name :score1.txt
The winners are :
['CU05', 'CU01', 'CU03']
```

```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO

def test(inputs,expected_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
```

<!-- grade_id: cell-318efae9bca2069c, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''top 3'''
test(['score1.txt'],"The winners are :\n['CU05', 'CU01', 'CU03']\n")
test(['score2.txt'],"The winners are :\n['CU07', 'CU05', 'CU01']\n")
test(['score3.txt'],"The winners are :\n['CU07', 'CU01', 'CU05']\n")
print('OK')
```

Output:
```
OK
```

<!-- grade_id: cell-c8599a43c8e1b8d8, points: 1, grade: true, solution: false -->
```python
'''เทสลับ1'''
```