## Q1 (8 points)
จงเขียนโปรแกรมที่อ่านเวกเตอร์ 2 เวกเตอร์จากไฟล์ vectors\<n\>.txt แล้วตรวจสอบว่ามีขนาดเท่ากันหรือไม่  <br>
- ถ้าไม่เท่ากัน จะแสดงข้อความว่า Incompatible size
- ถ้าเท่ากัน จะหา dot product ของเวกเตอร์ทั้งสอง <br>

สิ่งที่ต้องทำคือ

- สร้างฟังก์ชัน is_equal(v1,v2) เพื่อตรวจสอบความยาวของเวกเตอร์ทั้งสองเท่ากันหรือไม่ คืนค่าเป็น True หรือ False
- สร้างฟังก์ชัน dot(v1,v2) เพื่อหา dot product ของเวกเตอร์ทั้งสอง คืนค่าเป็น ผลรวมของ dot product
- สร้างฟังก์ชัน convert_to_float(v) เพื่อแปลงประเภทข้อมูลให้เป็น float คืนค่าเป็น v1 และ v2 ที่เป็นลิสต์
- สร้างฟังก์ชัน read_file_vectors(filename) เพื่ออ่านไฟล์และจัดรูปแบบเป็นลิสต์ คืนค่าเป็น v1 และ v2 ที่เป็นลิสต์

และแสดงผลลัพธ์ดังตัวอย่างข้างล่างนี้

### ตัวอย่างการทำงาน: 

ไฟล์ vectors1.txt:  <br>
6 4 9 6 -4 7 <br>
3 7 -3 9 -5 1 <br>
---
Choose your vector file: **vectors1.txt** <br>
v1 = [6.0, 4.0, 9.0, 6.0, -4.0, 7.0] <br>
v2 = [3.0, 7.0, -3.0, 9.0, -5.0, 1.0] <br>
v1*v2 = 100.0

### ตัวอย่างการทำงาน: 

ไฟล์ vectors2.txt: <br>
6 4 9 6 -4 7 <br>
3 7 -3 <br>
---
Choose your vector file: **vectors2.txt** <br>
v1 = [6.0, 4.0, 9.0, 6.0, -4.0, 7.0] <br>
v2 = [3.0, 7.0, -3.0] <br>
Incompatible size

### ตัวอย่างการทำงาน: 

ไฟล์ vectors3.txt: <br>
5.5 4.5 9.5 6 -4 7 <br>
3 7.1 -3.5 5 6 9 <br>

Choose your vector file: **vectors3.txt** <br>
v1 = [5.5, 4.5, 9.5, 6.0, -4.0, 7.0] <br>
v2 = [3.0, 7.1, -3.5, 5.0, 6.0, 9.0] <br>
v1*v2 = 84.2

### ตัวอย่างการทำงาน: 

ไฟล์ vectors4.txt: <br>
5.5 4 -9 6 -4 <br>
3 -7.1 0  <br>

Choose your vector file: **vectors4.txt** <br>
v1 = [5.5, 4.0, -9.0, 6.0, -4.0] <br>
v2 = [3.0, -7.1, 0.0] <br>
Incompatible size <br>

<!-- grade_id: name, points: 0, grade: false, solution: true -->
```python
def is_equal(v1,v2):
    if len(v1) == len(v2) :
        return True
    else :
        return False
def dot(v1,v2):
    b = 0
    for i in range (len(v1)) :
        a = v1[i]*v2[i]
        b = a+b
    return b
def convert_to_float(v):
    return [float(x) for x in v]
def read_file_vectors(filename):
    f = open(filename,'r') 
    lines = f.read().splitlines()
    v1 = lines[0].split()
    v2 = lines[1].split()
    f.close()
    return (v1,v2)
def run():
    x = input('Choose your vector file: ')
    y = read_file_vectors(x)
    v1 = convert_to_float(y[0])
    v2 = convert_to_float(y[1])
    print( "v1 =",v1 )
    print( "v2 =",v2 )
    if is_equal(v1,v2) :
        print('v1*v2 =',dot(v1,v2))
    else :
        print('Incompatible size')
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
Choose your vector file: vector4.txt
v1 = [5.5, 4.0, -9.0, 6.0, -4.0]
v2 = [3.0, -7.1, 0.0]
Incompatible size
```

```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO

def test(inputs,expected_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()
```

<!-- grade_id: cell-011068f4667a47f7, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบฟังก์ชัน is_equal'''
assert is_equal([6.0, 4.0, 9.0, 6.0, -4.0, 7.0],[3.0, 7.0, -3.0, 9.0, -5.0, 1.0]) == True
assert is_equal([1.0,2.0,3.0],[4.0,-5.1,-99.2]) == True
assert is_equal([1.0,2.0],[4.0,-5.1,-99.2]) == False
print('OK')
```

Output:
```
OK
```

<!-- grade_id: cell-231eb21054e1c1a8, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบฟังก์ชัน dot'''
assert dot([6.0, 4.0, 9.0, 6.0, -4.0, 7.0],[3.0, 7.0, -3.0, 9.0, -5.0, 1.0]) == 100.0
assert dot([1.0,2.0,3.0],[4.0,-5.1,-99.2]) == -303.8
print('OK')
```

Output:
```
OK
```

<!-- grade_id: cell-8552f0d30ed86b38, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบฟังก์ชัน convert_to_float'''
assert convert_to_float(['1','2','3']) == [1.0,2.0,3.0]
assert convert_to_float(['-55','-81','0.01']) == [-55.0, -81.0, 0.01]
print('OK')
```

Output:
```
OK
```

<!-- grade_id: cell-85b6a73599b8dcea, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบฟังก์ชัน read_file_vectors'''
assert read_file_vectors('vector2.txt') == (['6', '4', '9', '6', '-4', '7'], ['3', '7', '-3'])
assert read_file_vectors('vector3.txt') == (['5.5', '4.5', '9.5', '6', '-4', '7'], ['3', '7.1', '-3.5', '5', '6', '9'])
print('OK')
```

Output:
```
OK
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Vector dot product'''
test(['vector1.txt'],'v1 = [6.0, 4.0, 9.0, 6.0, -4.0, 7.0]\nv2 = [3.0, 7.0, -3.0, 9.0, -5.0, 1.0]\nv1*v2 = 100.0\n')
test(['vector3.txt'],'v1 = [5.5, 4.5, 9.5, 6.0, -4.0, 7.0]\nv2 = [3.0, 7.1, -3.5, 5.0, 6.0, 9.0]\nv1*v2 = 84.2\n')
print('OK')
```

Output:
```
OK
```

<!-- grade_id: cell-cbc47179f27fefee, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
'''Incompatible size'''
test(['vector2.txt'],'v1 = [6.0, 4.0, 9.0, 6.0, -4.0, 7.0]\nv2 = [3.0, 7.0, -3.0]\nIncompatible size\n')
test(['vector4.txt'],'v1 = [5.5, 4.0, -9.0, 6.0, -4.0]\nv2 = [3.0, -7.1, 0.0]\nIncompatible size\n')
print('OK')
```

Output:
```
OK
```

<!-- grade_id: cell-bae5e23d18255008, points: 1, grade: true, solution: false -->
```python
'''เทสลับ1'''
```

<!-- grade_id: cell-c5c71f7cd97f1236, points: 1, grade: true, solution: false -->
```python
'''เทสลับ2'''
```