## Q3 (4 points)
จงเขียนโปรแกรมที่อ่านไฟล์ studentScore.txt ที่มีรายชื่อนักเรียนและคะแนนสอบดังแสดงในตัวอย่างข้างล่างเก็บเป็น dictionary แล้วแสดงชื่อนักเรียนเรียงจากคะแนนสูงสุดไปต่ำสุด กำหนดให้ ไม่มีใครได้คะแนนเท่ากัน

```python

```

### ตัวอย่างไฟล์ studentScore1.txt: 

Piti 25 <br>
Mana 32 <br>
Ari 19 <br>
Suda 28 <br>
Somsak 26 <br>

### ตัวอย่างไฟล์ studentScore2.txt: 

Manee 22 <br>
Piti 32 <br>
Ton 35 <br>
Tree 21 <br>
Siri 23 <br>
Mon 20 <br>

### ตัวอย่างการทำงาน: 
Choose your studentScore: **studentScore1.txt** <br>
Mana <br>
Suda <br>
Somsak <br>
Piti <br>
Ari <br>

### ตัวอย่างการทำงาน: 
Choose your studentScore: **studentScore2.txt** <br>
Ton <br>
Piti <br>
Siri <br>
Manee <br>
Tree <br>

<!-- grade_id: name, points: 0, grade: false, solution: true -->
```python
def get_dict(file_name):
    st = open(file_name, 'r')
    sdict = {}    
    for line in st :
        name,score = line.split()
        score = int(score)
        sdict[name] = score
    st.close()
    return sdict
    
def run():
    p = input('Choose your studentScore: ')
    sdict = get_dict(p)
    scoredict = []
    for val in sdict.values():
        scoredict.append(val)
    scoredict.sort(reverse = True)

    list =[]
    for i in scoredict:     
        for j in sdict :
            if i == sdict[j] and j not in list :
                list.append(j)
    for q in list :
        print(q)           
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
run()
```

Output:
```
Choose your studentScore: studentScore3.txt
Kim
Ton
Piti
Siri
Manee
Tree
Mon
Lisa
```

```python
'''ตรวจสอบผลการคำนวณ'''
from unittest.mock import patch
from io import StringIO

def test(inputs,expected_output): 
    with patch('builtins.input', side_effect=inputs), patch('sys.stdout', new_callable=StringIO) as mock_stdout:
        run()
        assert mock_stdout.getvalue() == expected_output,mock_stdout.getvalue()       
```

<!-- grade_id: cell-159cd7c85a0c7a54, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
assert get_dict('studentScore1.txt') == {'Piti': 25, 'Mana': 32, 'Ari': 19, 'Suda': 28, 'Somsak': 26}
assert get_dict('studentScore2.txt') == {'Manee': 22, 'Piti': 32, 'Ton': 35, 'Tree': 21, 'Siri': 23, 'Mon': 20}
print('OK')
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
test(['studentScore1.txt'],'Mana\nSuda\nSomsak\nPiti\nAri\n') 
test(['studentScore2.txt'],'Ton\nPiti\nSiri\nManee\nTree\nMon\n') 
print('OK')
```

<!-- grade_id: cell-9241545eb2b30f20, points: 1, grade: true, solution: false -->
```python
'เทสลับ'
### BEGIN HIDDEN TESTS
assert get_dict('studentScore3.txt') == {'Manee': 22, 'Piti': 32, 'Ton': 35, 'Tree': 21, 'Siri': 23, 'Mon': 20, 'Kim': 40, 'Lisa': 18}
test(['studentScore3.txt'],'Kim\nTon\nPiti\nSiri\nManee\nTree\nMon\nLisa\n') 
### END HIDDEN TESTS  
```